# Introduction to Graph Theory in Sage

<b>SageMath</b> is an indispensable tool for the graph theorist.  In this notebook, we will learn how to use some of the built-in functions for graphs in Sage.

```sage
G = graphs.RandomGNP(8,0.5)  # A random graph on 8 vertices in which each edge is included with probability 0.5

# Feel free to change 8 to any positive integer you like, and 0.5 to any real number in [0,1].

G.show() # Display the graph G
```

What are each of the commands below doing?  Try regenerating the random graph above and then re-running the commands below to confirm your guesses.

```sage
G.order()
```

```sage
G.size()
```

```sage
G.vertices()
```

```sage
G.edges()
```

```sage
G.is_connected()
```

```sage
G.is_eulerian()
```

```sage
G.eulerian_circuit()
```

### Your Turn!

Write code below to determine whether or not the graph $G$ is Hamiltonian, and to find a Hamiltonian cycle in $G$, if it exists.

<b>Hint:</b> Search for "hamiltonian" in the <a href="https://doc.sagemath.org/html/en/reference/graphs/sage/graphs/generic_graph.html">SageMath documentation</a>.

```sage
# YOUR CODE HERE - Is G Hamiltonian?
```

```sage
# YOUR CODE HERE - If G is Hamiltonian, identify a Hamiltonian cycle
```

## Common Graphs

Do you want to know something about a specific graph?  There are several ways to define a specific graph in Sage.  First of all, there are many pre-defined families of graphs (paths, cycles, complete graphs, etc.).

```sage
P = graphs.PathGraph(4)  # A path on 4 vertices
P.show()
```

```sage
C = graphs.CycleGraph(5) # A cycle on 5 vertices
C.show()
```

```sage
B = graphs.CompleteBipartiteGraph(3,4) # A complete bipartite graph with partite sets of orders 3 and 4
B.show()
```

```sage
S = graphs.StarGraph(5) # A star with 5 leaves (hence 6 vertices total!)
S.show()
```

### Your Turn

Define and display a star graph on $8$ vertices. Assign this graph to the variable <code>S8</code>.

<!-- grade_id: cell-95a3ce16de7dba5f, points: 0, grade: false, solution: true -->
```sage
# YOUR CODE HERE
```

Does your graph <code>S8</code> have the correct number of vertices?

<!-- grade_id: cell-f724f037a1a51476, points: 1, grade: true, solution: false -->
```sage
"Check if S8 has 8 vertices. (1 mark)"
assert S8.order()==8, "S8 should have order 8."
print("Problem 1: Success!")
```

Define and display a complete graph on $5$ vertices.  If in doubt, search for "complete" in the <a href="https://doc.sagemath.org/html/en/reference/graphs/sage/graphs/graph_generators.html">SageMath documentation</a> for Common Graphs.

<!-- grade_id: cell-2a9a0aed56257bcf, points: 0, grade: false, solution: true -->
```sage
# YOUR CODE HERE
```

<!-- grade_id: cell-d86f8dfd2d89b6ea, points: 1, grade: true, solution: false -->
```sage
"Check if K5 has 5 vertices and 10 edges. (1 mark)"
assert K5.order()==5 and K5.size()==10, "K5 should have order 5 and size 10."
print("Problem 2: Success!")
```

## Custom Graphs

You can also define a graph using the <code>Graph</code> command.  You must provide a list of edges.

```sage
G = Graph() # Creates the null graph (with no vertices and no edges) and assigns it to the variable G
G.add_vertices([0,1,2,3,4]) # Add the vertices in the list [1,2,3,4,5] to G
G.add_edges([(0,1),(1,2),(2,3),(3,1)])
G.show()
```

Once a graph is constructed, you can add more vertices (or edges) whenever you like.  Note that <code>add_vertex</code> and <code>add_edge</code> take as input a single vertex or edge, respectively, while <code>add_vertices</code> and <code>add_edges</code> take as input a <it>list</it> of vertices or edges, respectively.

```sage
G.add_vertex(5)
G.add_edges([(3,4),(2,5)])
G.show()
```

### Your Turn

Construct a graph that is Eulerian but not Hamiltonian.  Assign it to the variable <code>E</code>.

<!-- grade_id: cell-89d30a43669b338c, points: 0, grade: false, solution: true -->
```sage
# YOUR CODE HERE
```

<!-- grade_id: cell-a1ec50fce8120213, points: 1, grade: true, solution: false -->
```sage
"Check if E is Eulerian. (1 mark)"
assert E.is_eulerian(), "E should be Eulerian."
print("Problem 3 Part 1: Success!")
```

<!-- grade_id: cell-3cb74d721619517b, points: 1, grade: true, solution: false -->
```sage
"Check if E is Hamiltonian. (1 mark)"
assert not E.is_hamiltonian(), "E should NOT be Hamiltonian."
print("Problem 3 Part 2: Success!")
```

Construct a graph that is Hamiltonian but not Eulerian.  Assign it to the variable <code>H</code>.

<!-- grade_id: cell-e8083fdba6e08517, points: 0, grade: false, solution: true -->
```sage
# YOUR CODE HERE
```

<!-- grade_id: cell-df92a344af3c129d, points: 1, grade: true, solution: false -->
```sage
"Check if H is Eulerian. (1 mark)"
assert not H.is_eulerian(), "H should NOT be Eulerian."
print("Problem 4 Part 1: Success!")
```

<!-- grade_id: cell-a3dd26b0a3b2d871, points: 1, grade: true, solution: false -->
```sage
"Check if H is Hamiltonian. (1 mark)"
assert H.is_hamiltonian(), "H should be Hamiltonian."
print("Problem 4 Part 2: Success!")
```